# Lab 6: Regular Expressions, SQL

## Course Policies

Here are some important course policies. These are also located at
http://www.ds100.org/sp18/.

**Collaboration Policy**

Data science is a collaborative activity. While you may talk with others about
the homework, we ask that you **write your solutions individually**. If you do
discuss the assignments with others please **include their names** at the top
of your solution.


## Due Date

This assignment is due at **11:59pm Monday, February 26th**. Instructions for submission are on the website.

```python
import pandas as pd
import numpy as np
import re
import sqlalchemy
```

## Objectives for Lab 6:

This lab has two main parts. 

In the first part, you will practice the basic usage of regular expressions and also learn to use `re` module in Python.  Some of the materials are based on the tutorial at http://opim.wharton.upenn.edu/~sok/idtresources/python/regex.pdf. As you work through the first part of the lab, you may also find the website http://regex101.com helpful. 

In the second part of the lab, we are going to practice viewing, sorting, grouping, and merging tables with SQL. 

---
# Part 1: Regular Expressions

We'll start by learning about the simplest possible regular expressions. Since regular expressions are used
to operate on strings, we'll start with the most common task: matching characters.

Most letters and characters will simply match themselves. For example, the regular expression `r'test'` will match the string `test` exactly. There are exceptions to this rule; some characters are special, and don't match themselves.

Here is a list of metacharacters that are widely used in regular experssion. 

<table border="1" class="dataframe" >
<thead>
  <tr style="text-align: right;">
    <th>Pattern </th>
    <th>Description</th> 
  </tr>
 </thead>
 <tbody>
  <tr>
    <td>^</td>
    <td>Matches beginning of line.</td> 
  </tr>
  <tr>
    <td>$</td>
    <td>Matches end of line.</td> 
  </tr>
  <tr>
    <td>.</td>
    <td>Matches any single character except newline. </td> 
  </tr>
  <tr>
    <td>*</td>
    <td>Matches 0 or more occurrences of preceding expression.</td>
  </tr>
  <tr>
    <td>+</td>
    <td>Matches 1 or more occurrence of preceding expression.</td>
  </tr>
  <tr>
    <td>?</td>
    <td>Matches 0 or 1 occurrence of preceding expression.</td>
  </tr>
  <tr>
    <td>[...]</td>
    <td>Matches any single character in brackets.</td>
  </tr>
  <tr>
    <td>[^...]</td>
    <td>Matches any single character <b>not</b> in brackets.</td>
  </tr>
  <tr>
    <td>{n}</td>
    <td>Matches exactly n number of occurrences of preceding expression.</td>
  </tr>
  <tr>
    <td>{n,}</td>
    <td>Matches n or more occurrences of preceding expression.</td>
  </tr>
  <tr>
    <td>{n,m}</td>
    <td>Matches at least n and at most m occurrences of preceding expression.</td>
  </tr>
  <tr>
    <td>a|b</td>
    <td>Matches either a or b.</td>
  </tr>
  <tr>
    <td>\1...\9</td>
    <td>Matches n-th grouped subexpression.</td>
  </tr>
  </tbody>
</table>


Perhaps the most important metacharacter is the backslash, ‘\’. As in Python string literals, the backslash
can be followed by various characters to signal various special sequences. It’s also used to escape all the
metacharacters so you can still match them in patterns; for example, if you need to match a `[` or `\`, you
can precede them with a backslash to remove their special meaning:  `\[` or `\\`. 

The following predefined special sequences are available:

<table border="1" class="dataframe" >
<thead>
  <tr style="text-align: right;">
    <th>Pattern </th>
    <th>Description</th> 
  </tr>
 </thead>
 <tbody>
  <tr>
    <td>\d</td>
    <td>Matches any decimal digit; this is equivalent to the class `[0-9]`</td> 
  </tr>
  <tr>
    <td>\D</td>
    <td>Matches any non-digit character; this is equivalent to the class `[^0-9]`.</td> 
  </tr>
  <tr>
    <td>\s</td>
    <td>Matches any whitespace character; this is equivalent to the class `[ \t\n\r\f\v]` </td> 
  </tr>
  <tr>
    <td>\S</td>
    <td>Matches any non-whitespace character; this is equivalent to the class `[^ \t\n\r\f\v]`.</td>
  </tr>
  <tr>
    <td>\w</td>
    <td>Matches any alphanumeric character; this is equivalent to the class `[a-zA-Z0-9_]`</td>
  </tr>
  <tr>
    <td>\W</td>
    <td>Matches any non-alphanumeric character; this is equivalent to the class `[^a-zA-Z0-9_]`.</td>
  </tr>
  </tbody>
</table>

---
<br/>
# Question 1
In this question, write patterns that match the given sequences. It may be as simple as the common letters on each line.

---
## Question 1a

Write a single regular expression to match the following strings without using the `|` operator.

1. **Match:** `abcdefg`
1. **Match:** `abcde`
1. **Match:** `abc`
1. **Skip:** `c abc`


<!-- grade_id: q1a-answer, points: 0, grade: false, solution: true -->
```python
regx1 = r"" # fill in your pattern
### BEGIN SOLUTION
regx1 = r"^abc[\w]*"
### END SOLUTION
```

<!-- grade_id: q1a-tests, points: 1, grade: true, solution: false -->
```python
assert ("|" not in regx1)
assert (re.search(regx1, "abc").group() == "abc") 
assert (re.search(regx1, "abcde").group() == "abcde") 
assert (re.search(regx1, "abcdefg").group() == "abcdefg")
assert (re.search(regx1, "c abc") is None)
```

---
## Question 1b

Write a single regular expression to match the following strings without using the `|` operator.

1. **Match:** `can`
1. **Match:** `man`
1. **Match:** `fan`
1. **Skip:** `dan`
1. **Skip:** `ran`
1. **Skip:** `pan`

<!-- grade_id: q1b-answer, points: 0, grade: false, solution: true -->
```python
regx2 = r"" # fill in your pattern
### BEGIN SOLUTION
regx2 = r"^([cmf]an)"
### END SOLUTION
```

<!-- grade_id: q1b-tests, points: 1, grade: true, solution: false -->
```python
assert ("|" not in regx2)
assert (re.match(regx2, 'can').group() == "can") 
assert (re.match(regx2, 'fan').group() == "fan") 
assert (re.match(regx2, 'man').group() == "man") 
assert (re.match(regx2, 'dan') is None) 
assert (re.match(regx2, 'ran') is None) 
assert (re.match(regx2, 'pan') is None)
```

---
<br/>
# Question 2

Now that we have written a few regular expressions, we are now ready to move beyond matching. In this question, we'll take a look at some methods from the `re` package.

---
## Question 2a:

Write a Python program to extract and print the numbers of a given string. 

1. **Hint:** use `re.findall`
2. **Hint:** use `\d` for digits and one of either `*` or `+`.

<!-- grade_id: q2a-answer, points: 0, grade: false, solution: true -->
```python
text_q2a = "Ten 10, Twenty 20, Thirty 30"

res_q2a = ...
### BEGIN SOLUTION
res_q2a = re.findall(r"\d+", text_q2a)
### END SOLUTION

res_q2a
```

<!-- grade_id: q2a-tests, points: 1, grade: true, solution: false -->
```python
assert res_q2a == ['10', '20', '30']
```

---
## Question 2b:

Write a Python program to replace at most 2 occurrences of space, comma, or dot with a colon.

**Hint:** use `re.sub(regex, "newtext", string, number_of_occurences)`


<!-- grade_id: q2b-answer, points: 0, grade: false, solution: true -->
```python
text_q2b = 'Python Exercises, PHP exercises.'
res_q2b = ... # Hint: use re.sub()
### BEGIN SOLUTION
res_q2b = re.sub("[ ,.]", ":", text_q2b, 2)
### END SOLUTION

res_q2b
```

<!-- grade_id: q2b-tests, points: 1, grade: true, solution: false -->
```python
assert res_q2b == 'Python:Exercises: PHP exercises.'
```

---
## Question 2c: 

Write a Python program to extract values between quotation marks of a string

<!-- grade_id: q2c-answer, points: 0, grade: false, solution: true -->
```python
text_q2c = '"Python", "PHP", "Java"'
res_q2c = ... # Hint: use re.findall()
### BEGIN SOLUTION
res_q2c = re.findall(r'"(.*?)"', text_q2c)
### END SOLUTION

res_q2c
```

<!-- grade_id: q2c-tests, points: 1, grade: true, solution: false -->
```python
assert res_q2c == ['Python', 'PHP', 'Java']
```

---
<br/>
# Part 2: SQL

In this part, we'll get some practice writing queries and manipulating tables. First, you'll need to connect to a hosted database in order to access the data. 

The cells below are copied from Biye's 2/22 lecture. Here, the `get_db_url` function will pick a random database each time you connect to try and balance the load across machines. Run the cell below to define the function.

```python
def get_db_url(password="dataisgreat"):
    import pandas as pd
    # Download the (potentially changing) list of databases
    db_list = pd.read_csv("http://ds100.org/sp18/assets/sql/db_list.csv")
    # Pick a random databse from the list
    db_info = db_list.sample().iloc[0]
    # Get it's info
    username = db_info['username']
    server = db_info['host']
    dbname = db_info['dbname']
    # Construct the database connection string
    db_url = 'postgres+psycopg2://{}:{}@{}/{}'.format(
        username, password, server, dbname)
    return db_url
```

Now, run the cells below to retrieve a db_url and connect to the hosted database.

```python
db_url = get_db_url()
db_url
```

```python
engine = sqlalchemy.create_engine(db_url, connect_args={'sslmode':'require'})
connection = engine.connect()
engine.table_names()
```

Out of the tables listed above, we'll use only two of them for questions 3-6:

+ `students`: Student name and other information; indexed by primary key `studentid`.
+ `ds100grades`: Grades per assignment per student. Can be linked to `students` using a _foreign key_ `studentid`.

Let's first take a peek at the students table.

---
## Question 3

Fill in the SQL query in `q3` to display all the data in the `students` table. (Don't worry, the table isn't big enough to crash your browser.)

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
q3 = '''
SELECT ...
FROM ...
'''
### BEGIN SOLUTION
q3 = '''
SELECT *
FROM students
'''
### END SOLUTION

pd.read_sql(q3, engine)
```

<!-- grade_id: q3-tests, points: 1, grade: true, solution: false -->
```python
assert len(connection.execute(q3).fetchall()) == 9
assert connection.execute(q3).fetchall()[0] == ('Sergey Brin', 2.8, 40, 'CS', 'M', 0)
```

---
## Question 4

Let's get a sense of how people are doing overall in this offering of DS100. Compute the average grade of all assignments in the `ds100grades` table in the `q4` string.

<!-- grade_id: q4-answer, points: 0, grade: false, solution: true -->
```python
q4 = '''
SELECT ...
FROM ...
'''
### BEGIN SOLUTION
q4 = '''
SELECT AVG(grade)
FROM ds100grades
'''
### END SOLUTION

pd.read_sql(q4, engine)
```

<!-- grade_id: q4-tests, points: 1, grade: true, solution: false -->
```python
assert len(connection.execute(q4).fetchall()) == 1
assert np.allclose(connection.execute(q4).fetchall()[0][0], 88.8730)
```

---
## Question 5

As you have likely realized, we would rather know the overall grade for each student. This will require us to join tables together. Write a query `q5` that joins the `student` table with the `ds100grades` table so that students are matched with their assignment grades.

You will find that each student appears 7 times (why?) for a total of 63 rows.

<!-- grade_id: q5-answer, points: 0, grade: false, solution: true -->
```python
# If you'd like, feel free to use a JOIN clause instead of the WHERE clause.
q5 = '''
SELECT ...
FROM ...
WHERE ...
'''
### BEGIN SOLUTION
q5 = '''
SELECT *
FROM students, ds100grades
WHERE students.studentid = ds100grades.studentid
'''
### END SOLUTION

pd.read_sql(q5, engine)
```

<!-- grade_id: q5-tests, points: 1, grade: true, solution: false -->
```python
assert len(connection.execute(q5).fetchall()) == 63
expected = [('Bill Gates', 1.0, 60, 'CS', 'M', 2, 3, 'hw', 2, 91.0), ('Danah Boyd', 3.9, 35, 'CS', 'F', 1, 29, 'labs', 1, 100.0), ('Hillary Mason', 4.0, 35, 'DATASCI', 'F', 3, 58, 'final', 3, 78.0), ('Mark Zuckerberg', 4.0, 30, 'CS', 'M', 5, 24, 'vitamins', 5, 96.0), ('Mike Olson', 3.7, 50, 'CS', 'M', 4, 50, 'midterm', 4, 73.0), ('Sheryl Sandberg', 4.0, 47, 'BUSINESS', 'F', 8, 18, 'proj', 8, 100.0), ('Susan Wojcicki', 4.0, 46, 'BUSINESS', 'F', 6, 43, 'participation', 6, 83.0)]
assert sorted(connection.execute(q5).fetchall())[::10] == expected
```

---
## Question 6

Now, write a query `q6` that calculates the average grade for each student assuming that all assignment types are weighted equally. (We know they aren't in reality, but we'll assume so for now)

Your query should output 9 rows, one for each student. Your table should contain two columns, the `name` of the student and the `avg` grade they got in DS100. Sort the result by descending order of average grade.

<!-- grade_id: q6-answer, points: 0, grade: false, solution: true -->
```python
q6 = '''
SELECT ...
FROM ...
WHERE ...
GROUP BY ...
ORDER BY ...
'''
### BEGIN SOLUTION
q6 = '''
SELECT s.name, AVG(g.grade)
FROM students AS s, ds100grades AS g
WHERE s.studentid = g.studentid
GROUP BY s.name
ORDER BY AVG(g.grade) DESC
'''
### END SOLUTION

pd.read_sql(q6, engine)
```

<!-- grade_id: q6-tests, points: 1, grade: true, solution: false -->
```python
assert len(connection.execute(q6).fetchall()) == 9
names = ['Sergey Brin', 'Danah Boyd', 'Marissa Meyer', 'Mike Olson', 'Susan Wojcicki', 'Bill Gates', 'Sheryl Sandberg', 'Hillary Mason', 'Mark Zuckerberg']
assert [name for name, grade in connection.execute(q6).fetchall()] == names
grades = [93.5714285714286, 91.4285714285714, 89.8571428571429, 89.0, 88.8571428571429, 87.1428571428571, 87.0, 86.8571428571429, 86.1428571428571]
assert np.allclose([grade for name, grade in connection.execute(q6).fetchall()], grades)
```

**Congrats! You are finished with this assignment. Don't forget to validate & submit before 11:59PM!**